// intro2.rs
// Make the code print a greeting to the world.
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE

fn main() {
    // to fix this one, we need to add a variable
    // containing a string input inside the brackets
    let text = "crustaceans";
    println!("Hello {text}!");
}
