// variables1.rs
// Make me compile!
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE

fn main() {
    // for this one to work 
    // we need to use the "let" statement 
    // to correctly define the variable
    let x = 5;
    println!("x has the value {}", x);
}
