// variables2.rs
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE

fn main() {
    // we need to define the type of the input and assign it a value
    let x: i32 = 0; // defining the type and then the value
    if x == 10 {
        println!("x is ten!");
    } else {
        println!("x is not ten!");
    }
}
